```python
# Initialize Otter
import otter
grader = otter.Notebook("hw02.ipynb")
```

# Homework 2: Python Modules `re` and `numpy`

Welcome to DATA 271: Data Wrangling and Visualization! This is your second homework assignment of the semester. 

In this assignment, we will be practice using some common and powerful Python modules, the `re` module and the `numpy` module. These will help you prepare for the rest of the semester. As always, if you get stuck or have questions feel free to discuss problems with your peers and come to office hours. 

For this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `my_list` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you passed previously. **NOTE:** This homework assignment contains hidden tests. Passing all the auto-graded tests in the assignment does *NOT* mean you have answered everything correctly. So be careful and check your work!

```python
import re
import numpy as np
import matplotlib.pyplot as plt
from csv import reader
```

## 1. Regex on Book Reviews

**Question 1.1.** You are analyzing book reviews to help decide your next read. Each review includes a star rating, the number of likes, and the number of comments, formatted as follows:
```python
"A beautifully written novel with unforgettable characters. Stars: 5 ⭐, Likes: 124, Comments: 37"
```

Write a regular expression that extracts the statistics (including stars, likes, and comments) from `review` below. Your answer should be a string. For example, for the review above, the answer would be:

```python
'Stars: 5 ⭐, Likes: 124, Comments: 37'
```

Please assign your regular expression in the `stats_regex` variable and assign the output statistics in the `stats` variable. 
**(5 Points)**

```python
review = "Interesting premise but slow execution. Stars: 3 ⭐, Likes: 56, Comments: 12"
stats_regex = ...
stats = ...
stats
```

```python
grader.check("q1_1")
```

**Question 1.2.** Write a regex that finds the star rating in `review` above. Your answer should be a string. For example, 
```python
'5'
```
Please assign your regular expression in the `stars_regex` variable and assign the output statistics in the `stars` variable.
**(5 Points)**

```python
stars_regex = ...
stars = ...
stars
```

```python
grader.check("q1_2")
```

**Question 1.3.** A sentiment analysis is a technique used to determine the emotional tone or attitude expressed in a piece of text. At its most basic level, sentiment analysis tries to classify whether a text is positive, negative, or neutral by looking for specific words and patterns. Think of it like a simple emotional detector - just as you might guess someone's mood by the words they use, sentiment analysis does this automatically with text.

In its simplest form, sentiment analysis involves looking for positive words (like "good," "happy," "excellent") and negative words (like "bad," "sad," "terrible") to make a judgment about the overall sentiment. While real-world sentiment analysis can be quite sophisticated using machine learning and natural language processing, we can start to understand the concept by using simple pattern matching to count positive and negative words in text.

Run the cell below to import book reviews for Mysterious Island by Jules Verne ([source](https://www.kaggle.com/datasets/meetnagadia/amazon-kindle-book-review-for-sentiment-analysis)). Use regular expression to perform a simple sentiment analysis on `book_reviews`. Assign `mysterious_island` to the 1 if the reviews indicate an overall attitude toward the book and assign 0 if the overall sentiment is negative. 

(*Hint* You will want to come up with your own list of positive and negative words. I have listed examples, but you should use of more. )
**(5 Points)**

```python
fhand = open('mysterious_island_reviews.txt', 'r') # file handle, open in read mode
book_reviews = fhand.read()
fhand.close()
```

```python

```

```python
mysterious_island = ...
```

```python
grader.check("q1_3")
```

<!-- BEGIN QUESTION -->

**Question 1.4.** Use regular expression and `book_reviews` to approximate how many reviewers gave Mysterious Island a poor rating (e.g. 1-2 stars). Assign `poor_ratings` to the number of poor ratings you expect, and explain your approach in text. 
**(5 Points)**

```python

```

```python
poor_ratings = ...
```

```python
grader.check("q1_4")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1.5.** Use regular expression and `book_reviews` to approximate how many reviewers gave Mysterious Island a positive rating (e.g. 1-2 stars). Assign `positive_ratings` to the number of poor ratings you expect and explain your approach in text. 
**(5 Points)**

```python

```

```python
positive_ratings = ...
positive_ratings
```

```python
grader.check("q1_5")
```

<!-- END QUESTION -->

## 2. Regex on dates

Pretend you are working with an environmental organization monitoring water quality off the coast of California. Your ultimate goal is to observe how water quality changes through time. The data you've been given are collected by voluteer citizens who go out and take samples of the water occasionally, but unfortunately all of the volunteers record dates in different styles. This makes it difficult to analyze the data through time, so you must first clean the data. 

Run the following cell to import the data. (These data were fabricated for the purpose of this exercise). 

```python
# Open and read the datasett
water = open('water_data_dates.csv')
water = reader(water) # type is csv.reader
water = [i[0] for i in list(water)]
water[0]=re.sub("\ufeff","",water[0])
water
```

**Question 2.1.** Fill in the regular expression to find all the long-style dates in `water`. For example, find all the dates in the form `January 2 2023` or `January 02 2023` or `Jan 2 2023` or `Jan 02 2023.` or `Jan 2nd 2023.` (The specific month, day, and year do not have to match these examples, just the form). 
**(5 Points)**

```python
regex = ...
long_dates = [re.findall(regex,date)[0] for date in water if re.findall(regex,date) !=[]]
long_dates
```

```python
grader.check("q2_1")
```

**Question 2.2.** Fill in the regular expression to find all the short-style dates in `water`. For example, find all the dates in the form `1/2/23` or `01/02/2023` or `01/02/23` or `1/2/2023` or `1-1-2023` or `01-01-2023` or `01-01-23`  or `1-1-23`. (The specific month, day, and year do not have to match these examples, just the form). 
**(5 Points)**

```python
regex2 = ...
short_dates = [re.findall(regex2,date)[0] for date in water if re.findall(regex2,date) !=[]] 
short_dates
```

```python
grader.check("q2_2")
```

**Question 2.3.** Start cleaning the dates in `long_dates` to put them all in the form of `MON D YEAR`. For example, `January 2 2023` should become `Jan 2 2023` and `January 21st 2023` should become `Jan 21 2023`. (Again, the specific month, day, and year do not have to match this example, just the form). Your final answer should be a list which is the same length as `long_dates`.
**(5 Points)**

```python
clean_long_dates = ...
clean_long_dates
```

```python
grader.check("q2_3")
```

## 3. Citations styles
Pretend you're working on a paper for your English class. Just hours before the deadline, you discover a citation style mix-up. You originally formatted your paper using Harvard style citations, thinking that was the required format. However, the assignment instructions actually asked for MLA style.

Your paper is long, with lots of carefully placed citations. Manually rewriting each citation would consume precious time, and you're worried you wouldn't finish before the deadline. You decide to use your regex skills to your advantage.

Run the cell below to import the text of your paper. *NOTE:* This text of the paper was generated with AI. It is not a genuine academic paper and should just be used practice regex techniques for education purposes.

```python
# Run this cell to import the paper
fhand = open('english_paper.txt', 'r') # file handle, open in read mode
english_paper = fhand.read()
fhand.close()
```

**Question 3.1** Make a regular expression that will match all the in-text citations. For example, it should match `(Smith et al. 2018)`, `(Johnson and Lee 2019)`, `(Taylor 2022)`, etc. Use your regular expression to find all the in-text citations in the English paper, and put them in a list called `in_text_matches`. **(5 points)**

```python
in_text_regex = ...
in_text_matches = ...
in_text_matches
```

```python
grader.check("q3_1")
```

**Question 3.2** One difference between Harvard and MLA citation styles is that MLA requires a comma between the authors and the year in in-text citations. For example, (Smith et al. 2018) should be changed to (Smith et al., 2018). Use regular expression to make these changes. Update the paper's text and store it in `english_paper_update1`. **(5 points)**

```python
regex = ...
english_paper_update1 = ...
english_paper_update1
```

```python
grader.check("q3_2")
```

**Question 3.3** Another difference between Harvard and MLA citation styles is that MLA uses an ampersand symbol (&) instead of the word "and" when there are two authors. For example, (Sharma and Gupta, 2024) should be changed to (Sharma & Gupta, 2024). Starting with `english_paper_update1`, use regular expression to make these changes and save them in `english_paper_update2`. **(5 points)**

```python
english_paper_update2 = ...
english_paper_update2
```

```python
grader.check("q3_3")
```

## 4. Numpy for Financial Analysis
Below is the daily open data for the S&P 500 in January 2025 ([source](https://www.investing.com/indices/us-spx-500-historical-data)).

```python
# Stock prices
stock_prices = np.array(
[6096.79,
 6050.75,
 6057.7,
 6026.97,
 5969.04,
 6121.43,
 6076.32,
 6081.39,
 6014.12,
 5995.4,
 5963.61,
 5905.21,
 5859.27,
 5782.02,
 5890.35,
 5910.66,
 5993.26,
 5982.81,
 5891.07,
 5903.26])
```

**Question 4.1.** Calculate the daily percent change in stock prices and find the standard deviation of the percent change. Recall that percent change is 

$$\frac{\text{new\_value-old\_value}}{\text{old\_value}}\times100$$

*HINT:* You might find `np.diff` helpful. 
**(5 Points)**

```python
# Calculate daily percentage change
percent_change = ...

# Find standard deviation
sd_change = ...
sd_change
```

```python
grader.check("q4_1")
```

**Question 4.2:** *Trend analysis:* A moving average is a statistical technique used to smooth out fluctuations or noise in a dataset, making it easier to identify trends or patterns over time. It calculates the average of a subset of consecutive data points within a specified window or interval, and then moves the window forward one data point at a time, recalculating the average each time.

Compute the 5-day moving average of stock prices to identify long-term trends in stock prices. Put your results in a numpy array called `moving_average`.

*Hint:* To compute a moving average:

- Set the window size: 5 in this case
- Calculate the average: Take the average of the data points within the window at each step. This average becomes the value of the moving average at that position.
- Move the window: Slide the window forward by one data point and recalculate the average. Repeat this process until the window reaches the end of the dataset.
 
**(5 Points)**

```python
window_size = ...
moving_average = ...

...
    moving_average = ...

moving_average
```

```python
grader.check("q4_2")
```

**Question 4.3:** Now recompute the percent change from your moving average and find the standard deviation. What do you observe? **(5 Points)**

```python
# Calculate percentage change
percent_change_ma = ...

# Find standard deviation
sd_change_ma = ...
sd_change_ma
```

```python
grader.check("q4_3")
```

## 5. Numpy for Student Performance
Suppose you have scores for students across multiple exams in a class. Each row represents a student, and each column represents an exam. (These data are fabricated for the purpose of this exercise. We will generate the score randomly in the cell below).

```python
num_students = 50
num_exams = 4
exam_scores = np.random.normal(82, 7, size=(num_students, num_exams))
```

**Question 5.1:** Compute the average score for each exams. Your answer should be an array with each entry being the average score of each exam. 
**(5 Points)**

```python
average_scores = ...
average_scores
```

```python
grader.check("q5_1")
```

**Question 5.2:** The first three columns are midterms worth 20% each, and the last column is the final worth 40% of the overall grade. Compute each student's overall grade in the class. 
**(5 Points)**

```python
student_grades = ...
student_grades
```

```python
grader.check("q5_2")
```

**Question 5.3:** How did the student who got the highest grade overall do on exam 2?
**(5 Points)**

```python
exam_2_grade = ...
exam_2_grade
```

```python
grader.check("q5_3")
```

## 6. Random Functions in Numpy for Simulation: Estimating $\pi$

In this problem we will use Numpy's random number generation to estimate $\pi$.

Consider the unit circle inscribed in a square. The circle, with radius $1$ has an area of $A = \pi$.  The square with sides of length $2$ has an area of $4$. The ratio of the area of the circle to the area of the square is therefore $\frac{\pi}{4}$.  If we throw darts randomly, the number of darts that land  in the circle divided by the total number of darts thrown will be described by the ratio of $\frac{\pi}{4}$.

It is enough to just consider the quarter circle and quarter square in the first quadrant, because the ratio in one quadrant should be the same as the ratio for the whole inscribed circle. Because the equation of the unit circle is $x^2+y^2 = 1$ we know a point is within the quarter circle if $\sqrt{x^2 + y^2} <1 $. Otherwise, it falls outside of the circle.

**Question 6.1:** Generate an array of size $1000 \times 2$ with random numbers between 0 and 1. The first column will represent the x-coordinates of 1000 randomly thrown darts, and the second column will be the y-coordinates. **(5 Points)**

```python
random_darts = ...
```

```python
grader.check("q6_1")
```

**Question 6.2:** Create an array with all the points from `random_darts` that fall inside the first quadrant of the unit circle. **(5 Points)**

```python
inside = ...
inside
```

```python
grader.check("q6_2")
```

**Question 6.3:** Use your previous two answers to get an approximation of $\pi$. **(5 Points)**

```python
pi_approx = ...
pi_approx
```

```python
grader.check("q6_3")
```

Run the following cell to visualize where your random darts landed. The blue ones are inside the unit circle, the red ones are outside. 

```python
# To visualize
plt.figure(figsize = (8, 8))
plt.scatter(random_darts[:, 0], random_darts[:, 1], s= 1, c = 'red') 
plt.scatter(inside[:, 0], inside[:, 1], s = 1, c = 'blue')
plt.show()
```

## You're done!

Congratulations! You finished Homework 2! Download the zip file and upload to Canvas. 

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1_1 (points: [0])
```
>>> stats != ... and stats_regex != ...
True
```


Tests for q1_2 (points: [0])
```
>>> stars != ... and stars_regex != ...
True
```


Tests for q1_3 (points: [0])
```
>>> mysterious_island != ...
True
```


Tests for q1_4 (points: None)
```
>>> poor_ratings != ...
True
```


Tests for q1_4_manual (points: [])


Tests for q1_5 (points: None)
```
>>> positive_ratings != ...
True
```


Tests for q1_5_manual (points: [])


Tests for q2_1 (points: [0])
```
>>> regex != ...
True
```


Tests for q2_2 (points: [0])
```
>>> regex2 != ...
True
```


Tests for q2_3 (points: [0])
```
>>> clean_long_dates != ...
True
```


Tests for q3_1 (points: [0])
```
>>> in_text_matches != ...
True
```


Tests for q3_2 (points: [0])
```
>>> english_paper_update1 != ...
True
```


Tests for q3_3 (points: [0])
```
>>> english_paper_update2 != ...
True
```


Tests for q4_1 (points: [0])
```
>>> len(percent_change) == 19
True
```


Tests for q4_2 (points: [0, 0])
```
>>> window_size != ...
True
```
```
>>> len(moving_average)
16
```


Tests for q4_3 (points: [0])
```
>>> len(percent_change_ma)
15
```


Tests for q5_1 (points: [0])
```
>>> len(average_scores)
4
```


Tests for q5_2 (points: [0])
```
>>> len(student_grades)
50
```


Tests for q5_3 (points: [0])
```
>>> exam_2_grade != ...
True
```


Tests for q6_1 (points: [0])
```
>>> random_darts.dtype == np.array([0.1]).dtype
True
```


Tests for q6_2 (points: [0])
```
>>> 700 <= len(inside) <= 850
True
```


Tests for q6_3 (points: [0])
```
>>> 3 < pi_approx < 3.3
True
```
