```python
# Initialize Otter
import otter
grader = otter.Notebook("hw06.ipynb")
```

# Homework 6

Welcome to Homework 6 of DATA 271: Data Wrangling and Visualization! In this assignment, we will practice scraping data from a website and practice combining and cleaning data. As always, if you get stuck or have questions feel free to discuss problems with your peers and come to office hours. 

For this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `my_list` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you passed previously. **NOTE:** This homework assignment contains hidden tests. Passing all the auto-graded tests in the assignment does *NOT* mean you have answered everything correctly. So be careful and check your work!

### Import libraries

**Question 0:** Import the appropriate modules for this assignment including NumPy, Pandas, Matplotlib, and Seaborn, Regular Expression (re), Requests, and BeautifulSoup. **(5 Points)**

```python
import sys # keep this import
```

### 1. Importing the data

In this assignment, we will explore data from the men's soccer team at Cal Poly Humboldt. Let's look at the [2023](https://humboldtathletics.com/sports/mens-soccer/stats/2023) season.  

**Question 1.1:** Import the data from the link above. **(5 Points)**

```python
stats_page = ...
```

```python
grader.check("q1_1")
```

**Question 1.2:** Check the status code to make sure your request was successful. **(5 Points)**

```python
stats_page_status = ...
stats_page_status
```

```python
grader.check("q1_2")
```

**Question 1.3:** Inspect the text of your requests object. **(5 Points)**

```python
stats_text = ...
stats_text
```

```python
grader.check("q1_3")
```

### 2. Parse the data

**Question 2.1:** Use BeautifulSoup to begin parsing this HTML data. Create a BeautifulSoup object from `stats_text`. **(5 Points)**

```python
soup = ...
soup
```

```python
grader.check("q2_1")
```

**Question 2.2:** Start using your soup object to explore the contents of the soccer stats webpage. How many level 1 headers tables are on the webpage? **(5 Points)**

```python
num_h1_headers = ...
num_h1_headers
```

```python
grader.check("q2_2")
```

**Question 2.3:** What is the text content of the second h1 header? **(5 Points)**

```python
h1_header_text = ...
h1_header_text
```

```python
grader.check("q2_3")
```

**Question 2.4:** Find the number of tables on the webpage.
**(5 Points)**

```python
num_tables = ...
num_tables
```

```python
grader.check("q2_4")
```

**Question 2.5:** Inspect the contents of a few of the tables. Based on your exploration, what statistics are in the sixth table on the webpage? 

1. Game-By-Game Offensive Stats Statistics
2. Game-By-Game Results Statistics
3. Individual Overall Goalkeeping Statistics
4. Team Leaders: Game Winning Goals

Enter 1, 2, 3, or 4 in `table6_content`.
**(5 Points)**

```python
# Inspect the tables
```

```python
table6_content = ...
```

```python
grader.check("q2_5")
```

**Question 2.6:** Create a Pandas DataFrame from the sixth table on the webpage.

*HINT:* Notice that the column headers have a slightly different structure in HTML. You might find it helpful to first create the data frame so that it contains the body of the table. Then create the column names of the table afterward. Make sure that the `Date` column is a Pandas datetime type. **(5 Points)**

```python
table_data = ...


stats_df = ...
stats_df.columns = ...
stats_df
```

```python
grader.check("q2_6")
```

**Question 2.7:** The seventh table on the webpage contains Humboldt's offensive stats. Create a Pandas DataFrame from the the seventh table. Remove the `Total` and `Opponent` rows.

*HINT:* You will likely have to use the same method you used in the previous problem for making the data frame column names. **(5 Points)**

```python
table_data = ...

        
humboldt_offense_df = ...
humboldt_offense_df.columns = ...
humboldt_offense_df 
```

```python
grader.check("q2_7")
```

**Question 2.8:** The ninth table on the webpage contains Humboldt's opponents' offensive stats. Create a Pandas DataFrame from the ninth table. 

*HINT:* You will likely have to use the same method you used in the previous problem for making the data frame column names. **(5 Points)**
**(5 Points)**

```python
table_data = ...

        

opponent_offense_df = ...
opponent_offense_df.columns = ...
opponent_offense_df 
```

```python
grader.check("q2_8")
```

### 3. Clean/Preprocess the Data

**Question 3.1:** The offense datasets have the following columns:
- Date: Represents the date of the game.
- Opponent: Indicates the team or opponent Cal Poly Humboldt played against.
- Score: Shows the final score of the game, in the format "home team score - away team score".
- G: Represents the number of goals scored by the team.
- A: Represents the number of assists made by players on the team.
- PTS: Stands for "Points" which is a combination of goals and assists. Goals are worth 2 points and assists are worth 1 point. 
- SH: Represents the total number of shots taken by the team.
- Shot%: Represents the shooting percentage, calculated as goals divided by shots taken.
- SOG: Stands for "Shots on Goal" and represents the number of shots that were on target and required a save from the goalkeeper.
- SOG%: Represents the shots on goal percentage, calculated as shots on goal divided by total shots taken.
- YC-RC: Represents Yellow Cards and Red Cards. Yellow cards are typically given as warnings for fouls, and red cards result in ejection from the game. This column might be formatted as "YC-RC" to indicate both types of cards.
- GW: Stands for "Game-Winning" and represents goals or assists that directly led to winning the game.
- PK-ATT: Stands for "Penalty Kick Attempts" and represents the number of penalty kicks attempted by the team.
- Min: Represents the minutes played.

Several of these columns won't be useful for our analysis, and others should be reformatted. In `humboldt_offense_df`, split the `YC-RC` column into two separate columns. One should be called `YC` and contain the number of yellow cards Humboldt recieved. The other new column should be called `RC` and contain the number of red cards Humboldt recieved. Once done, drop the `SOG`,`SOG%`,`YC-RC`,`GW`,`PK-ATT`,`Min` columns from `humboldt_offense_df`.
**(5 Points)**

```python
humboldt_offense_df['YC'] = ...
humboldt_offense_df['RC'] = ...
humboldt_offense_df
```

```python
grader.check("q3_1")
```

**Question 3.2:** Perform a similar task for `opponent_offense_df`. Split the `YC-RC` column into two separate columns. One should be called `YC` and contain the number of yellow cards the opponent recieved. The other new column should be called `RC` and contain the number of red cards the opponent recieved. Once done, drop the `Opponent`,`Score`,`SOG`,`SOG%`,`YC-RC`,`GW`,`PK-ATT`,`Min` columns from `opponent_offense_df`. **(5 Points)**

```python
opponent_offense_df['YC'] = ...
opponent_offense_df['RC'] = ...
opponent_offense_df
```

```python
grader.check("q3_2")
```

**Question 3.3:** Check the dtypes of `humboldt_offense_df`. Convert the data types of each column appropriately. **(5 Points)**

```python
# Check dtypes
```

```python
# Convert dtypes
humboldt_offense_df
```

```python
grader.check("q3_3")
```

**Question 3.4:** Perform a similar task for `opponent_offense_df`. Convert the data types of each column appropriately. **(5 Points)**

```python
# Check dtypes
```

```python
# Convert dtypes
opponent_offense_df.dtypes
```

```python
grader.check("q3_4")
```

**Question 3.5:** Look at Humboldt's offensive statistics with their opponents offensive statistics. Combine your `humboldt_offense_df` and `opponent_offense_df` dataframes into a single data frame. 

*HINT*: The combined data should have additional columns. Some of those columns will show Humboldt's offensive stats, and other columns will show the opponent's offensive stats. Give the columns corresponding to Humboldt's stats the suffix `_humboldt` and the columns corresponding to opponents the suffix `_opp`  **(5 Points)**

```python
combined_offense_stats = humboldt_offense_df.merge(opponent_offense_df, on = 'Date', suffixes=['_humboldt', '_opp'])
combined_offense_stats
```

```python
grader.check("q3_5")
```

**Question 3.6:** In this analysis, we will want to compare some of the offensive statistics to Humboldt's win/loss record. Add the `W/L` column from `stats_df` to your `combined_offense_stats` data frame appropriately. **(5 Points)**

```python
all_offense_stats = ...
all_offense_stats
```

```python
grader.check("q3_6")
```

### 4. Visualize the data

**Question 4.1:** Create a countplot showing the number of games Humboldt won, lost, and tied in the 2023 season. Comment on what you see. **(5 Points)**

```python
wins_plot = ...
plt.show()
```

```python
grader.check("q4_1")
```

**Question 4.2:** Create a scatter plot showing the relationship between Humboldt's Shot% and their opponents' Shot%. Map the color of points to the W/L variable. Comment on what you notice. **(5 Points)**

```python
shot_compare_plot = sns.scatterplot(data = all_offense_stats, x = 'Shot%_humboldt', y = 'Shot%_opp', hue = 'W/L')
plt.show()
```

```python
grader.check("q4_2")
```

**Question 4.3:** Create a Seaborn violin plot showing the distribution and 5 number summary of Shot% for Humboldt and their opponent in each category of `W/L`. In other words, there should be a a total of 6 violins in your plot. There should be 2 violins for each category of `W/L`. One should show Humboldt's Shot% distribution (and be one color), and the other should show their opponents' Shot% distribution (and be another color). Comment on the visualization.

*HINT:* You might find it helpful to reshape the data using the `.melt` method and use `hue` to show Humboldt vs. the opponents. **(5 Points)**

```python
melted_offense = ...
melted_offense
```

```python
win_loss_shot_violins = ...
```

```python
grader.check("q4_3")
```

<!-- BEGIN QUESTION -->

**Question 4.4:** Create two more plots to explore this dataset. Feel free to reshape/combine data as needed or parse additional data from the HTML. **(5 Points)**

_Type your answer here, replacing this text._

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q0 (points: [])


Tests for q1_1 (points: [0])
```
>>> bool(stats_page != ...)
True
```


Tests for q1_2 (points: [0])
```
>>> bool(stats_page_status != ...)
True
```


Tests for q1_3 (points: [0])
```
>>> bool(stats_text != ...)
True
```


Tests for q2_1 (points: [0])
```
>>> bool(soup != ...)
True
```


Tests for q2_2 (points: [0])
```
>>> bool(num_h1_headers != ...)
True
```


Tests for q2_3 (points: [0, 2])
```
>>> bool(h1_header_text != ...)
True
```
```
>>> str(h1_header_text)
"2023 Men's Soccer Cumulative Statistics"
```


Tests for q2_4 (points: [0])
```
>>> bool(num_tables != ...)
True
```


Tests for q2_5 (points: [0])
```
>>> bool(table6_content != ...)
True
```


Tests for q2_6 (points: [0])
```
>>> bool(type(stats_df) is pd.DataFrame)
True
```


Tests for q2_7 (points: [0])
```
>>> bool(type(humboldt_offense_df) is pd.DataFrame)
True
```


Tests for q2_8 (points: [0])
```
>>> bool(type(opponent_offense_df) is pd.DataFrame)
True
```


Tests for q3_1 (points: [0])
```
>>> humboldt_offense_df.shape
(17, 10)
```


Tests for q3_2 (points: [0])
```
>>> opponent_offense_df.shape
(17, 8)
```


Tests for q3_3 (points: [0])
```
>>> bool(humboldt_offense_df.dtypes.G == type(1.0) or humboldt_offense_df.dtypes.G == type(1))
True
```


Tests for q3_4 (points: [0])
```
>>> bool(opponent_offense_df.dtypes.G == type(1.0) or opponent_offense_df.dtypes.G == type(1))
True
```


Tests for q3_5 (points: [0])
```
>>> bool('Date' in combined_offense_stats.columns)
True
```


Tests for q3_6 (points: [0])
```
>>> bool('W/L' in all_offense_stats)
True
```


Tests for q4_1 (points: [0])
```
>>> bool(wins_plot != ...)
True
```


Tests for q4_2 (points: [0])
```
>>> bool(shot_compare_plot != ...)
True
```


Tests for q4_3 (points: [0])
```
>>> bool(win_loss_shot_violins != ...)
True
```


Tests for q4_4_manual (points: [])
