```python
# Initialize Otter
import otter
grader = otter.Notebook("hw04.ipynb")
```

# Homework 4: Exploring and Visualizing Data

Welcome to homework 4 of DATA 271: Data Wrangling and Visualization! In this assignment, we will practice exploring data with visualizations and summaries. As always, if you get stuck or have questions feel free to discuss problems with your peers and come to office hours. 

For this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `my_list` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you passed previously. **NOTE:** This homework assignment contains hidden tests. Passing all the auto-graded tests in the assignment does *NOT* mean you have answered everything correctly. So be careful and check your work!

<!-- BEGIN QUESTION -->

**Question 0:** Import the appropriate modules for this assignment including NumPy, Pandas, Matplotlib, and Seaborn. **(5 Points)**

```python

```

<!-- END QUESTION -->

## 1. Grammar of Graphics

<!-- BEGIN QUESTION -->

**Question 1.1.** What are the main components of a layered grammar of graphics? List them below and give a short description of each. **(5 Points)**

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1.2.** Based on what you see in the figure below, deconstruct the visualization into it's grammar of graphics components. 

<img src="cars_graphic.PNG" alt="drawing" width="500"/>

**(5 Points)**

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1.3.** Based on what you see in the figure below, deconstruct the visualization into it's grammar of graphics components. 

<img src="country_graphic.PNG" alt="drawing" width="500"/>

**(5 Points)**

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# 2. Exploritory Data Analysis on Penguins
In the following problem we will be working with the `penguins.csv` file (already imported into your local repository if you are on JupyterHub). 

The data were collected and made available by Dr. Kristen Gorman and the Palmer Station, Antarctica LTER, a member of the Long Term Ecological Research Network. 

We will work through and exploritory data analysis on this dataset, and maybe learn a little about penguins along the way!

**Question 2.1:** Import the `penguins.csv` as a Pandas DataFrame. **(5 Points)**

```python
penguins = ...
penguins.head()
```

```python
grader.check("q2_1")
```

**Question 2.2:** How many penguins are in the dataset? (i.e. How many rows are there?) **(5 Points)**

```python
num_penguins = ...
num_penguins
```

```python
grader.check("q2_2")
```

**Question 2.3:** Create a Python list of the column labels in your `penguins` DataFrame. **(5 Points)**

```python
col_names = ...
col_names
```

```python
grader.check("q2_3")
```

In case you aren't familiar with some of these terms, feel free to check out [this paper](https://doi.org/10.1371/journal.pone.0090081) where the data were originally published to learn more. Here is a diagram of the bill length and depth too.

<img src="culmen_depth.PNG" alt="drawing" width="400"/>

Artwork by @allison_horst

**Question 2.4:** How many species of penguins are in the dataset? **(5 Points)**

```python
num_species = ...
num_species
```

```python
grader.check("q2_4")
```

**Question 2.5:** How many islands are in the dataset? **(5 Points)**

```python
num_islands = ...
num_islands
```

```python
grader.check("q2_5")
```

**Question 2.6:** Are there any null values in the dataset? Assign `null_vals` to `True` if yes and `False` if no. **(5 Points)**

```python
null_vals = ...
null_vals
```

```python
grader.check("q2_6")
```

**Question 2.7:** Which column in the dataset has the largest mean? Which one has the smallest mean?

*NOTE:* Your answers should be strings of appropriate column names. **(5 Points)**

```python
max_mean = ...
min_mean = ...
max_mean, min_mean
```

```python
grader.check("q2_7")
```

**Question 2.8:** Create a Pandas Series that shows how many total penguins there are on each island. Then make a barplot to visualize your results. (Be sure to label your axes!) 

*HINT:* If you are using Seaborn, try using `seaborn.countplot`.  
**(5 Points)**

```python
# Create your series here
penguins_per_island = ...
penguins_per_island
```

```python
# Create your plot here
```

```python
grader.check("q2_8")
```

**Question 2.9:** Create a Pandas Series that shows how many total penguins there are of each species. Then make a barplot to visualize your results. (Be sure to label your axes!) 

*HINT:* If you are using Seaborn, try using `seaborn.countplot`.  
**(5 Points)**

```python
# Create your series here
penguins_per_species = ...
penguins_per_species
```

```python
# Create your plot here
```

```python
grader.check("q2_9")
```

**Question 2.10.1:** Create a Pandas Series that shows how many total penguins there are of each species on each island. Then make a barplot to visualize your results. (Be sure to label your axes!) 

*HINT:* If you are using Seaborn, try using `seaborn.countplot`.  
**(5 Points)**

```python
species_per_island = ...
species_per_island
```

```python

```

```python
grader.check("q2_10_1")
```

**Question 2.10.2:** Use your solutions from the problem above to answer the following True or False questions. 

- **Statement 1:** Adelie penguins are on all the islands in the dataset. 
- **Statement 2:** Gentoo penguins are only on Biscoe island in the dataset.
- **Statement 3:** All species are on Dream island in the dataset.
- **Statement 4:** There are more Gentoo penguins than Adelie penguins on Biscoe island.
- **Statement 5:** Chinstrap penguins are only on Torgenson island in the dataset.  

Assign each statement to a bool indicating whether it is true or false. 
**(5 Points)**

```python
statement1 = ...
statement2 = ...
statement3 = ...
statement4 = ...
statement5 = ...
```

```python
grader.check("q2_10_2")
```

<!-- BEGIN QUESTION -->

**Question 2.11:** Let's start exploring the numeric variables. Use seaborn to look at the distributions of and relationships between all the numeric variables. Comment on anything interesting you find.
**(2 Points)**

_Type your answer here, replacing this text._

```python
# Create your visulization here
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2.12:** Some of the pairplots above have interesting gaps. Use seaborn to make a pairplot of all the numeric variables where each species is a different color. Comment on anything interesting you find.
**(5 Points)**

_Type your answer here, replacing this text._

```python
# Create your visulization here
```

<!-- END QUESTION -->

**Question 2.12:** How do the numeric variables correlate with one another? Use a heatmap to visualize the correlations and then answer the following true or false values. 

- **Statement A:** Bill length and Bill depth are positively correlated. 
- **Statement B:** Body mass and flipper length are positively correlated. 
- **Statement C:** Body mass and bill depth are negatively correlated. 
  
Assign each statement to a bool indicating whether it is true or false.

**(5 Points)**

```python
# Create your visulization here
plt.show()
```

```python
statement_a = ...
statement_b = ...
statement_c = ...
```

```python
grader.check("q2_13")
```

**Question 2.14:** 

**Scenario 1:** Suppose you found a penguin that had a flipper length of 240 mm and bill depth of 15 mm, what species would you guess the penguin is?  
**Scenario 2:** Suppose you found a penguin that had a bill depth of 18 mm and a body mass of 3500 g, what species would you guess the penguin is?  
**Scenario 3:** Suppose you found a penguin that had a flipper length of 190 mm and a bill length of 37, what species would you guess the penguin is?

1. Adelie
2. Chinstrap
3. Gentoo
4. Too difficult to determine

Enter 1, 2, 3, 4 for the scenarios below. 
**(3 Points)**

```python
scenario1 = ...
scenario2 = ...
scenario3 = ...
```

```python
grader.check("q2_14")
```

<!-- BEGIN QUESTION -->

**Question 2.15:** Based on your answers above, what attributes (or combinations of attributes) of penguins might be helpful for predicting penguin species. Which attributes (or combinations of attributes) would not be very helpful?
**(2 Points)**

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<img src="lter_penguins.PNG" alt="drawing" width="400"/>
Artwork by @allison_horst

# 3. Exploritory Data Analysis on Diabetes

<!-- BEGIN QUESTION -->

**Question 3.1:** Import `diabetes_prediction_dataset.csv` ([source](https://www.kaggle.com/datasets/iammustafatz/diabetes-prediction-dataset/data)) and perform your own exploratory analysis on it. You can use the steps we went through on the penguins dataset as inspiration, but be creative and think if your own ways of exploring and visualizing the data too. Explain what you find with markdown cells as you go. 

Identify columns (if any) in the dataset that could potentially help predict if a patient has diabetes. 
**(10 Points)**

_Type your answer here, replacing this text._

```python
df = ...
df
```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

<!-- END QUESTION -->

## You're done! 
Congrats on finishing Homework 4! Run the following cell to download a zip and submit to Canvas. 

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q0_manual (points: [])


Tests for q1_1_manual (points: [])


Tests for q1_2_manual (points: [])


Tests for q1_3_manual (points: [])


Tests for q2_1 (points: [0])
```
>>> not penguins.empty
True
```


Tests for q2_10_1 (points: [0])
```
>>> (species_per_island != ...).sum()
np.int64(5)
```


Tests for q2_10_2 (points: [0])
```
>>> statement1 != ... and statement2 != ... and (statement3 != ...) and (statement4 != ...) and (statement5 != ...)
True
```


Tests for q2_11_manual (points: [])


Tests for q2_12_manual (points: [])


Tests for q2_13 (points: [0])
```
>>> statement_a != ... and statement_b != ... and (statement_c != ...)
True
```


Tests for q2_14 (points: [0])
```
>>> type(scenario1) == int and type(scenario2) == int and (type(scenario3) == int)
True
```


Tests for q2_2 (points: [0])
```
>>> num_penguins != ...
True
```


Tests for q2_3 (points: [0])
```
>>> col_names != ...
True
```


Tests for q2_4 (points: [0])
```
>>> num_species != ...
True
```


Tests for q2_5 (points: [0])
```
>>> num_islands != ...
True
```


Tests for q2_6 (points: [0])
```
>>> null_vals != ...
True
```


Tests for q2_7 (points: [0])
```
>>> max_mean != ... and min_mean != ...
True
```


Tests for q2_8 (points: [0])
```
>>> (penguins_per_island != ...).sum()
np.int64(3)
```


Tests for q2_9 (points: [0])
```
>>> (penguins_per_island != ...).sum()
np.int64(3)
```


Tests for q3_1_manual (points: [])
