```python
# Initialize Otter
import otter
grader = otter.Notebook("lab10.ipynb")
```

# Lab 10: Web Scraping and Manipulating Data

Welcome to Lab 10 of Data Wrangling and Visualization!

## Overview
The web provides a rich source of data, so we often want to work with that information programmatically in order to make sense of it. Sometimes, that data is provided to us by website creators via csv files or through an API (Application Programming Interface). Other times, we need to collect text from the web ourselves.

It is easy to pull HTML from a website but more difficult to find the information we want from HTML.  Parsing the HTML for targeted information and then storing that information in a structured format will be the focus of this activity. Once we have our data in a structured format, we will practice cleaning and reshaping the data.


## In today's lab, we will
- Become familiar with the structure of HTML and CSS by inspecting source code
- Use the Requests and Beautiful Soup libraries to acquire and parse data from websites
- Work through the challenges of web scraping (i.e. turning messy, unstructured data into workable form)
- Use Pandas methods to clean, manipulate, and reshape data for a brief EDA

This activity will make use of the Requests and Beautiful Soup Python modules.
Documentation can be found here: 
- https://docs.python-requests.org/en/latest/
- https://www.crummy.com/software/BeautifulSoup/

### Import packages


```python
import pandas as pd
import requests
from bs4 import BeautifulSoup
import warnings
warnings.filterwarnings("ignore")
import seaborn as sns
sns.set_style('darkgrid')
import matplotlib.pyplot as plt
import numpy as np
```

### 1. Scrape and Beautify Cal Poly Humboldt Data

**Question 1.1:** Go to https://pine.humboldt.edu/anstud/cgi-bin/filter.pl?relevant=ft1yr_retention.out. Read through the webpage and inspect its source code. What information does this page provide? Once you have a feel for what is on the webpage, import its contents with the requests module.

```python
# Import the data
page = ...
```

```python
grader.check("q1_1")
```

Output:
```
q1_1 results: All test cases passed!
```

**Question 1.2:** Inspect the text of the data you just imported. 

```python
# inspect the text
page_text = ...
page_text
```

Output:
```
'<HTML>\n<HEAD> <script src="http://www.humboldt.edu/anstud/ga_pine_irp.js" type="text/javascript"></script>  \n<TITLE> First-Time full-time Undergrad Retention 1-year cohorts</TITLE>\n</HEAD>\n<BODY #FFFFFF>\n<p><table border=1 cellpadding=5><TR BGCOLOR=#EEEEEE><th colspan=3>First-Time full-time Undergrads (Freshmen) Retention reports<tr BGCOLOR=#FFFFEE><td>Freshmen cohorts</td><td><a href="filter.pl?relevant=ft1term_retention.out">1-term retention</a>\n&nbsp &nbsp &nbsp<a href="filter.pl?relevant=ft1yr_retention.out">1-year retention</a>\n&nbsp &nbsp &nbsp<a href="filter.pl?relevant=ft2yr_retention.out">2-year retention</a>\n&nbsp &nbsp &nbsp<a href="filter.pl?relevant=ft3yr_retention.out">3-year retention</a>\n&nbsp &nbsp &nbsp</td></tr><tr BGCOLOR=#EEEEEE><td><a href="filter.pl?relevant=trans1yr_retention_UD.out">Transfer Student Retention reports</a>\n&nbsp &nbsp &nbsp</td><td><a href="https://ie.humboldt.edu/grad-reten"> Return to Graduation and Retention data</a></td></tr></tabl
```

```python
grader.check("q1_2")
```

Output:
```
q1_2 results: All test cases passed!
```

**Question 1.3:** Create a Beautiful Soup object from `page_text`. 

```python
# Create soup object
soup = ...
soup
```

Output:
```
<html>
<head> <script src="http://www.humboldt.edu/anstud/ga_pine_irp.js" type="text/javascript"></script>
<title> First-Time full-time Undergrad Retention 1-year cohorts</title>
</head>
<body #ffffff="">
<p></p><table border="1" cellpadding="5"><tr bgcolor="#EEEEEE"><th colspan="3">First-Time full-time Undergrads (Freshmen) Retention reports</th></tr><tr bgcolor="#FFFFEE"><td>Freshmen cohorts</td><td><a href="filter.pl?relevant=ft1term_retention.out">1-term retention</a>
     <a href="filter.pl?relevant=ft1yr_retention.out">1-year retention</a>
     <a href="filter.pl?relevant=ft2yr_retention.out">2-year retention</a>
     <a href="filter.pl?relevant=ft3yr_retention.out">3-year retention</a>
     </td></tr><tr bgcolor="#EEEEEE"><td><a href="filter.pl?relevant=trans1yr_retention_UD.out">Transfer Student Retention reports</a>
     </td><td><a href="https://ie.humboldt.edu/grad-reten"> Return to Graduation and Retention data</a></td></tr></table><center><h2>First-Time full-time Undergrad
```

```python
grader.check("q1_3")
```

Output:
```
q1_3 results: All test cases passed!
```

<!-- BEGIN QUESTION -->

**Question 1.4:** Use the `prettify()` method to turn the Beautiful Soup parse tree into a nicely formatted string with each HTML tag on its own line. Take a look through it and take note of any tags or attributes you are unfamiliear with. Look up any tags that aren't familiar.

```python
# Print pretty data


```

Output:
```
'<html>\n <head>\n  <script src="http://www.humboldt.edu/anstud/ga_pine_irp.js" type="text/javascript">\n  </script>\n  <title>\n   First-Time full-time Undergrad Retention 1-year cohorts\n  </title>\n </head>\n <body #ffffff="">\n  <p>\n  </p>\n  <table border="1" cellpadding="5">\n   <tr bgcolor="#EEEEEE">\n    <th colspan="3">\n     First-Time full-time Undergrads (Freshmen) Retention reports\n    </th>\n   </tr>\n   <tr bgcolor="#FFFFEE">\n    <td>\n     Freshmen cohorts\n    </td>\n    <td>\n     <a href="filter.pl?relevant=ft1term_retention.out">\n      1-term retention\n     </a>\n     <a href="filter.pl?relevant=ft1yr_retention.out">\n      1-year retention\n     </a>\n     <a href="filter.pl?relevant=ft2yr_retention.out">\n      2-year retention\n     </a>\n     <a href="filter.pl?relevant=ft3yr_retention.out">\n      3-year retention\n     </a>\n    </td>\n   </tr>\n   <tr bgcolor="#EEEEEE">\n    <td>\n     <a href="filter.pl?relevant=trans1yr_retention_UD.out">\n      Transf
```

<!-- END QUESTION -->

### 2. Parse and clean the data

**Question 2.1:** Use a BeautifulSoup method to determine how many tables are on the webpage. 

```python
num_tables = ...
num_tables
```

Output:
```
16
```

```python
grader.check("q2_1")
```

Output:
```
q2_1 results: All test cases passed!
```

**Question 2.2:** Access the second to last table and put the contents in `table1_code`. Inspect it, and then give a brief description of what the table contains. 

```python
table1_code = ....
table1_code
```

Output:
```
<table align="CENTER" bgcolor="#FFFFEE" border="1" cellpadding="1" cellspacing="0" width="100%">
<tr><th bgcolor="#EEEEEE" colspan="11">First-Time full-time Undergrads by current Major - cohort and pct return 1 year later  
<br/><small>ft1yr_retention report generated: 03-OCT-25</small></th></tr>
<tr bgcolor="#EEEEEE"><th>Major Program<br/></th><th>Fall 16<br/></th><th>Fall 17<br/></th><th>Fall 18<br/></th><th>Fall 19<br/></th><th>Fall 20<br/></th><th>Fall 21<br/></th><th>Fall 22<br/></th><th>Fall 23<br/></th><th>Fall 24<br/></th><th>Fall 25<br/></th></tr>
<tr align="RIGHT"><td align="LEFT">Anthropology</td><td>23</td><td>15</td><td>18</td><td>19</td><td>7</td><td>13</td><td>16</td><td>16</td><td>8</td><td>9</td></tr>
<tr align="RIGHT"><td align="LEFT"><br/></td><td>73.9%</td><td>86.7%</td><td>83.3%</td><td>57.9%</td><td>100.0%</td><td>76.9%</td><td>75.0%</td><td>87.5%</td><td>87.5%</td><td>.0%</td></tr>
<tr align="RIGHT"><td align="LEFT">Applied Fire Science &amp; Mgmt</td><td>0</td><
```

```python
grader.check("q2_2")
```

Output:
```
q2_2 results: All test cases passed!
```

**Question 2.3:** Put the contents from the second to last table (`table1_code`) into Pandas Data Frame.

*HINT 1:* You will probably find it helpful to first create a nested list of the data. Each inner list should contain the contents of one row in the table. The outer list should contain all the rows. Use this nested list to create the Pandas DataFrame.

*HINT 2:* Notice that the first two rows in the table have headers (a slightly different structure from the rest of the table). You might find it helpful to first create the data frame to contain the body of the table. Then create the column names with the headers afterward. That is, find the table headers with `.find_all('th')`

table1_df

```python
# Create a nested list of data

nested_list = []

.....

    
nested_list
```

Output:
```
[[],
 [],
 ['Anthropology', '23', '15', '18', '19', '7', '13', '16', '16', '8', '9'],
 ['',
  '73.9%',
  '86.7%',
  '83.3%',
  '57.9%',
  '100.0%',
  '76.9%',
  '75.0%',
  '87.5%',
  '87.5%',
  '.0%'],
 ['Applied Fire Science & Mgmt',
  '0',
  '0',
  '0',
  '0',
  '0',
  '2',
  '4',
  '16',
  '15',
  '35'],
 ['', '%', '%', '%', '%', '%', '50.0%', '100.0%', '56.3%', '73.3%', '.0%'],
 ['Art', '36', '36', '38', '17', '20', '33', '41', '41', '17', '13'],
 ['',
  '75.0%',
  '72.2%',
  '65.8%',
  '70.6%',
  '80.0%',
  '81.8%',
  '80.5%',
  '75.6%',
  '76.5%',
  '.0%'],
 ['Art History and Museum Studies',
  '0',
  '0',
  '0',
  '0',
  '0',
  '0',
  '0',
  '6',
  '3',
  '2'],
 ['', '%', '%', '%', '%', '%', '%', '%', '100.0%', '33.3%', '.0%'],
 ['Biochemistry', '0', '0', '1', '0', '0', '0', '4', '4', '18', '15'],
 ['', '%', '%', '100.0%', '%', '%', '%', '50.0%', '100.0%', '72.2%', '.0%'],
 ['Biology', '66', '80', '74', '61', '56', '48', '61', '41', '43', '52'],
 ['',
  '72.7%',
  '68.8%',
  '70
```

```python
# Put the data in a Pandas dataframe and rename columns as needed
table1_df = ....


table1_df
```

Output:
```
0                   Anthropology     23     15     18     19       7     13  \
0                   Anthropology     23     15     18     19       7     13   
1                                 73.9%  86.7%  83.3%  57.9%  100.0%  76.9%   
2    Applied Fire Science & Mgmt      0      0      0      0       0      2   
3                                     %      %      %      %       %  50.0%   
4                            Art     36     36     38     17      20     33   
..                           ...    ...    ...    ...    ...     ...    ...   
129                               39.3%  35.5%  38.0%  50.9%   33.3%  30.0%   
130                     Wildlife     29     53     51     51      41     41   
131                               82.8%  84.9%  78.4%  76.5%   87.8%  80.5%   
132                      Zoology     34     37     40     47      27     26   
133                               76.5%  62.2%  80.0%  63.8%   77.8%  69.2%   

0        16     16      8    9  
0        16     16
```

```python
# rename the columns

table1_df.columns = ...

table1_df
```

Output:
```
Major Program Fall 16 Fall 17 Fall 18 Fall 19 Fall 20  \
0                   Anthropology      23      15      18      19       7   
1                                  73.9%   86.7%   83.3%   57.9%  100.0%   
2    Applied Fire Science & Mgmt       0       0       0       0       0   
3                                      %       %       %       %       %   
4                            Art      36      36      38      17      20   
..                           ...     ...     ...     ...     ...     ...   
129                                39.3%   35.5%   38.0%   50.9%   33.3%   
130                     Wildlife      29      53      51      51      41   
131                                82.8%   84.9%   78.4%   76.5%   87.8%   
132                      Zoology      34      37      40      47      27   
133                                76.5%   62.2%   80.0%   63.8%   77.8%   

    Fall 21 Fall 22 Fall 23 Fall 24 Fall 25  
0        13      16      16       8       9  
1     76.9%   
```

```python
grader.check("q2_3")
```

**Question 2.4:** Notice that even rows contain numbers and odd rows contain percents in `table1_df1`. Let's separate this into two Pandas data frames. One should contain only numbers, one should contain percents. 

*NOTE:* Make sure that both tables have data in the `Major Program` column. Only the even rows have an entry for `Major Program`!

```python
# Create numbers data frame
table1_nums = ...
table1_nums
```

```python
# Create percents data frame
table1_percents = ...
table1_percents
```

```python
grader.check("q2_4")
```

<!-- BEGIN QUESTION -->

**Question 2.5:** Check the data types of each column in both tables from the previous question. 

```python
# check nums dtypes
```

```python
# check percents dtypes
```

<!-- END QUESTION -->

**Question 2.6:** Cast the datatypes in your tables appropriately. In other words, if a column contains numbers, cast it to an appropriate numeric data type.

*NOTE:* The percents in the `table1_percents` DataFrame should be floats.

```python
# Cast datatypes in nums table
    
table1_nums.dtypes
```

```python
# Cast datatypes in percents table
table1_percents.dtypes
```

```python
grader.check("q2_6")
```

## 3. Visualize

**Question 3.1:** Visualize the relationship between the 1st Year retention rates (percents) in Fall 14 and the 1st Year retention rates (percents) in Fall 23. Make sure your visualization has a title and your axes are labeled appropriately as Fall 14 and Fall 23.

```python
scatter_16_vs_23 = ...
plt.show()
```

```python
grader.check("q3_1")
```

<!-- BEGIN QUESTION -->

**Question 3.2:** Visualize the trend in 1st Year retention rates (percents) over time for Biology. On the same axes plot the trend for English (in a different color). Make sure your plot has an informative legend, axes labels, and title. 

*HINT:* You will probably find it helpful to reshape your `table1_percents` dataframe. First, use the transpose method: `df.T`...this will get you columns of Year and Biology and English. Then you will need to tidy up the dataframe a little: rename the columns, `.drop('Major Program')` and then reset the index. Lastly, you will use `melt()` to get a long form dataframe with the columns of year/term, Major Program with the levels `Biology` and `English`, and the percent.

```python
# reshape the data as needed

```

```python
# plot the trends
retention_trends = ...
plt.show()
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 3.3:** Create one more data visualization of your choice. This might require reshaping your data. Describe what insights you gain from your visualization.

```python
# Visualize
```

<!-- END QUESTION -->

## You're done! 

Congratulations on finishing the lab! Gus is proud of you! Run the cell below and submit to Canvas. 

<img src="gus_big_stretch.JPG" alt="drawing" width="500"/>

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1_1 (points: None)
```
>>> bool(page != ...)
True
```
```
>>> int(page.status_code)
200
```


Tests for q1_2 (points: None)
```
>>> bool(page_text != ...)
True
```
```
>>> bool(type(page_text) is str)
True
```


Tests for q1_3 (points: None)
```
>>> bool(soup != ...)
True
```
```
>>> bool(type(soup) is BeautifulSoup)
True
```


Tests for q1_4_manual (points: None)


Tests for q2_1 (points: None)
```
>>> bool(num_tables != ...)
True
```
```
>>> int(num_tables)
16
```


Tests for q2_2 (points: None)
```
>>> bool(table1_code != ...)
True
```
```
>>> bool('First-Time full-time Undergrads by current Major' in table1_code.find('tr').text)
True
```


Tests for q2_3 (points: None)
```
>>> bool(type(table1_df) is pd.DataFrame)
True
```
```
>>> int(sum([i in ['Major Program', 'Fall 14', 'Fall 14', 'Fall 15', 'Fall 16', 'Fall 17', 'Fall 18', 'Fall 19', 'Fall 20', 'Fall 21', 'Fall 22', 'Fall 23'] for i in list(table1_df.columns)]))
9
```
```
>>> int(table1_df.shape[0])
134
```


Tests for q2_4 (points: None)
```
>>> bool(type(table1_nums) is pd.DataFrame and type(table1_percents) is pd.DataFrame)
True
```
```
>>> int(sum([i in ['Major Program', 'Fall 14', 'Fall 14', 'Fall 15', 'Fall 16', 'Fall 17', 'Fall 18', 'Fall 19', 'Fall 20', 'Fall 21', 'Fall 22', 'Fall 23'] for i in list(table1_nums.columns)]))
9
```
```
>>> int(sum([i in ['Major Program', 'Fall 14', 'Fall 14', 'Fall 15', 'Fall 16', 'Fall 17', 'Fall 18', 'Fall 19', 'Fall 20', 'Fall 21', 'Fall 22', 'Fall 23'] for i in list(table1_percents.columns)]))
9
```
```
>>> table1_nums.shape[0]
67
```
```
>>> 'Art' in table1_percents['Major Program'].values
True
```


Tests for q2_5manual (points: None)


Tests for q2_6 (points: None)
```
>>> bool(table1_nums.iloc[:, 5].dtype is np.array(2).dtype)
True
```
```
>>> bool(table1_percents.iloc[:, 5].dtype is np.ones(2).dtype)
True
```


Tests for q3_1 (points: None)
```
>>> bool('Fall 16' == scatter_16_vs_23.get_xlabel() or 'Fall 16' == scatter_16_vs_23.get_ylabel())
True
```
```
>>> bool('Fall 23' == scatter_16_vs_23.get_xlabel() or 'Fall 23' == scatter_16_vs_23.get_ylabel())
True
```
```
>>> bool(len(scatter_16_vs_23.get_title()) != 0)
True
```


Tests for q3_2_manual (points: None)


Tests for q3_3_manual (points: None)
