```python
# Initialize Otter
import otter
grader = otter.Notebook("hw01.ipynb")
```

```python

```

# ASSIGNMENT CONFIG
init_cell: true 
export_cell: true 
requirements: ../../requirements.txt
export_cell:
    pdf: true
    force_save: false
solutions_pdf: false
template_pdf: false
generate:
    points_possible: 100.0
    show_stdout: true
    filtering: true
    pagebreaks: true
    zips: false

# Homework 1: Python Basics

Welcome to DATA 271: Data Wrangling and Visualization! This is your first homework assignment of the semester. 

In this assignment, we will be practicing Python basics to prepare for the rest of the semester. Some of these concepts may be review from DATA 111, others have been covered for the first time in DATA 271. As always, if you get stuck or have questions feel free to discuss problems with your peers and come to office hours. 

For this homework and all future ones, please be sure to not re-assign variables throughout the notebook. For example, if you use `my_list` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you passed previously. **NOTE:** This homework assignment contains hidden tests. Passing all the auto-graded tests in the assignment does *NOT* mean you have answered everything correctly. So be careful and check your work!

## 1. Variable assignments, Data Types, and Built-in Functions & Methods

<!-- BEGIN QUESTION -->

**Question 1.1.** Assign three variables `a`, `b`, and `c`. Assign values `10`, `20.5`, and `'30'` respectively. Calculate the sum of `a` and `b`, then concatenate `c` to the result. Print the final result. **(5 Points)**

```python
a = ...
b = ...
c = ...
a_b_sum = ...
final_result = ...
...
```

```python
grader.check("q1_1")
```

**Question 1.2.** Identify the error in the following variable assignment: 

`1st_variable = 10`

1. There is no error

2. Invalid variable name

3. Missing a semicolon

4. Incorrect indentation

Assign `problem` to 1, 2, 3, or 4. 
**(5 Points)**

```python
problem = ...
```

```python
grader.check("q1_2")
```

<!-- BEGIN QUESTION -->

**Question 1.3.** Explain why the following code would result in an error

```{python}
statement = 'Harry is 23 years old.'
age = statement[9:11]

if age > 18:
    print('They are an adult.')
```
**(5 Points)**

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 1.4.** Convert `pi_string` to an int.
**(5 Points)**

```python
pi_string = '3.14159'
pi_int = ...
pi_int
```

```python
grader.check("q1_4")
```

**Question 1.5.** Given the sentence below, use string or list methods to perform the following tasks:
- Count the number of words. 
- Find the longest word.
- Print the sentence with the words sorted alphabetically. (*HINT*: Your output should be a string, not a list).
**(5 Points)**

```python
sentence = 'data science is an fun class'

word_count = ...
print(word_count)

longest_word = ...
print(longest_word)

sorted_alphabetically = ...
print(sorted_alphabetically)
```

```python
grader.check("q1_5")
```

**Question 1.6.** Assign the variable `state` to `'California'`, `city` to `'Arcata'`, `school` to `'Cal Poly Humboldt'`, and `data_class` to `271` in *one line of code*. 
**(5 Points)**

```python

```

```python
grader.check("q1_6")
```

## 2. Functions

**Question 2.1.** Create a function called `power_sum` that takes two arguments, `x` and `n`, and returns the sum of the first `n` powers of `x`. *Note:* Don't include the zeroth power.
**(5 Points)**

```python
...
    current_sum = ...
    ...
        current_sum = ...
    ...
```

```python
grader.check("q2_1")
```

**Question 2.2.** Write a function called `is_palindrome` that takes a string as input and returns `True` if it is a palindrome, and `False` otherwise. Recall that a palindrome is a word, phrase, or sequence that reads the same backward as forward.
**(5 Points)**

```python
...
    ...
```

```python
grader.check("q2_2")
```

## 3. Conditional statements

**Question 3.1.** Identify the error in the following code: 

```{python}
num = 15
if num < 10:
    print("Less than 10")
elif num < 20
    print("Between 10 and 20")
```

1. The elif statement is missing a colon.

2. The elif condition is missing an operator.

3. The if statement is not indented correctly.

4. There is no error in the code.

Assign `problem2` to 1, 2, 3, or 4. 
**(5 Points)**

```python
problem2 = ...
```

```python
grader.check("q3_1")
```

**Question 3.2:** Write a function called `analyze_numbers` that takes two numbers as arguments. 
- If the first number is positive and the second number is negative, return "First positive, second negative" 
- If the first number is negative and the second number is positive, return "First negative, second positive" 
- If both are positive, return their product. 
- If both are negative, return their sum. 
- Return "One or both numbers are zero" if one or both numbers are zero. 
**(5 Points)**

```python
...
    ...
        ...
    ...
        ...
    ...
        ...
    ...
        ...
    ...
        ...
```

```python
grader.check("q3_2")
```

**Question 3.3:** Write a function `username(first, middle, last)` that returns an 8-character username based on a person’s first, middle, and last names.

The format for the username is all lower case: first letter of first name; first letter of middle name, and then up to 6 letters of the last name. Longer last names should be truncated after 6 characters. For shorter last names, fill in with `1`s so that there are 8 characters in the username. 
**(5 Points)**

```python
...
    ...
        ...
    ...
        ...
```

```python
grader.check("q3_3")
```

**Question 3.4:** Write a function called `cat_dog(n)` that accepts an integer as an argument and returns a list from 1 to `n` (inclusive), but with the following conditions:

- If a number is divisible by 4, the list element should be "Woof" instead of the number.
- If a number is divisible by 6, the list element should be "Meow" instead of the number.
- If a number is divisible by both 4 and 6 (i.e., divisible by 12), the list element should be "WoofMeow".
- Otherwise, the list element is the number itself.
**(5 Points)**

```python
...

```

```python
grader.check("q3_4")
```

## 4. Lists and Dictionaries

**Question 4.1:** Use list comprehension to create a list containing the cubes of even numbers from 1 to 10 (inclusive).
**(5 Points)**

```python
list_of_even_cubes = ...
list_of_even_cubes
```

```python
grader.check("q4_1")
```

**Question 4.2:** Create a list `orig_list` of integers from 1 to 10 (inclusive). Using list comprehension, create a new list `primes` containing only the prime numbers. Recall that a prime number is a whole number greater than 1 that cannot be exactly divided by any whole number other than itself and 1. 

*Hint:* You might find the `all()` function useful. It returns `True` if all items in an iterable are true, otherwise it returns False. See the example below
**(5 Points)**

```python
# Example of all() function
print(all([2 < x for x in range(2,10)]))
print(all([2 < x for x in range(3,10)]))
```

```python
orig_list = ...
primes = ...
primes
```

```python
grader.check("q4_2")
```

**Question 4.3:** Consider the dictionary `MLB_teams` below that maps a location to the name of its corresponding Major League Baseball team. Add the Texas Rangers and the Kansas City Royals to the dictionary. (*Hint*: Be sure to use "Kansas City" as the key, not "Kansas")
**(5 Points)**

```python
MLB_teams = {
    'Colorado' : 'Rockies',
    'Boston'   : 'Red Sox',
    'Minnesota': 'Twins',
    'Milwaukee': 'Brewers',
    'Seattle'  : 'Mariners',
    'St. Louis': 'Cardinals',
    'Los Angeles': 'Dodgers',
    'San Fransico': 'Giants',
    'San Diego': 'Padres'
}
```

```python
MLB_teams
```

```python
grader.check("q4_3")
```

**Question 4.4:** Continuing with the `MLB_teams` dictionary from the previous question, make a list of all of the locations sorted alphabetically by their team name. *NOTE:* Your final list should not have the team names in it, only locations.  
**(5 Points)**

```python
locations_sorted_by_team = ...
locations_sorted_by_team
```

```python
grader.check("q4_4")
```

**Question 4.5:** Given two lists, `keys` and `values`, create a dictionary where each key-value pair comes from corresponding elements in the lists.  
**(5 Points)**

```python
keys = ['egg','peanuts','tofu']
values = [6, 26, 36]

grams_of_protein = ...
grams_of_protein
```

```python
grader.check("q4_5")
```

**Question 4.6:** Use dictionary comprehension to create a dictionary where keys are characters from a string, and values are the count of occurrences in the string given. *NOTE:* Assume we want the counts to not be case sensitive. 
**(5 Points)**

```python
given_string = 'Fuzzy Wuzzy was a bear. Fuzzy Wuzzy had no hair. \
If Fuzzy Wuzzy had no hair, then Fuzzy Wuzzy wasn’t very fuzzy, was he?'

character_counts = ...
character_counts
```

```python
grader.check("q4_6")
```

**Question 4.7:** You have a dictionary representing presidential candidate statistics. Use dictionary methods to find the candidate with the highest average number of media mentions.
**(5 Points)**

(Note: Data was pulled in January 2024 from this [source](https://www.brookings.edu/articles/the-2024-presidential-primary-tracker/))

```python
candidates = {
    "Ron DeSantis": {"party": 'republican', "media mentions": 39.48},
    "Joe Biden": {"party": 'democrat', "media mentions": 87.80},
    "Donald Trump": {"party": 'republican', "media mentions": 113.20},
    "Nikki Haley": {"party": 'republican', "media mentions": 24.78},
    "Dean Phillips": {"party": 'democrat', "media mentions": 5.24}
}

most_mentioned_candidate = ...
most_mentioned_candidate
```

```python
grader.check("q4_7")
```

## 5. Reminders from DATA 111

**Question 5.1:** In Python, what is the purpose of the for loop?

1. To define a function

2. To make decisions based on conditions

3. To create an infinite loop

4. To iterate over elements in a sequence

Assign `for_purpose` to 1,2,3, or 4.
**(5 Points)**

```python
for_purpose = ...
```

```python
grader.check("q5_1")
```

**Question 5.2:** In Python, what is the purpose of the elif statement?

1. To catch exceptions

2. To define a new function

3. To execute a block of code if the previous conditions are False

4. To create an infinite loop

Assign `elif_purpose` to 1,2,3, or 4.
**(5 Points)**

```python
elif_purpose = ...
```

```python
grader.check("q5_2")
```

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1_1 (points: [0])
```
>>> a != ... and b != ... and (c != ...) and (a_b_sum != ...) and (final_result != ...)
True
```


Tests for q1_2 (points: [0])
```
>>> 1 <= problem <= 4
True
```


Tests for q1_3_manual (points: [])


Tests for q1_4 (points: [0])
```
>>> pi_int != ...
True
```


Tests for q1_5 (points: [0, 0])
```
>>> word_count != ... and longest_word != ... and (sorted_alphabetically != ...)
True
```
```
>>> type(sorted_alphabetically) == str
True
```


Tests for q1_6 (points: [0])
```
>>> type(state) == str and type(city) == str and (type(school) == str) and (type(data_class) == int)
True
```


Tests for q2_1 (points: [0])
```
>>> type(power_sum(2, 3)) == int or type(power_sum(2, 3)) == float
True
```


Tests for q2_2 (points: [0])
```
>>> type(is_palindrome('radar')) == bool
True
```


Tests for q3_1 (points: [0])
```
>>> problem2 != ...
True
```


Tests for q3_2 (points: [0])
```
>>> analyze_numbers(2, 3) != None
True
```


Tests for q3_3 (points: [0, 0, 0, 0])
```
>>> len(username('Michael', 'Gary', 'Scott')) == 8
True
```
```
>>> len(username('Pamela', 'Morgan', 'Beesly')) == 8
True
```
```
>>> len(username('Holly', 'Partridge', 'Flax')) == 8
True
```
```
>>> len(username('Dwight', 'Kurt', 'Schrute')) == 8
True
```


Tests for q3_4 (points: [0, 0])
```
>>> type(cat_dog(1)) == list
True
```
```
>>> len(cat_dog(10)) == 10
True
```


Tests for q4_1 (points: [0])
```
>>> list_of_even_cubes != ...
True
```


Tests for q4_2 (points: [0])
```
>>> primes != ...
True
```


Tests for q4_3 (points: [0])
```
>>> len(MLB_teams) == 11
True
```


Tests for q4_4 (points: [0])
```
>>> type(locations_sorted_by_team) == list
True
```


Tests for q4_5 (points: [0])
```
>>> type(grams_of_protein) == dict
True
```


Tests for q4_6 (points: [0])
```
>>> type(character_counts) == dict
True
```


Tests for q4_7 (points: [0])
```
>>> most_mentioned_candidate != ...
True
```


Tests for q5_1 (points: [0])
```
>>> 1 <= for_purpose <= 4
True
```


Tests for q5_2 (points: [0])
```
>>> 1 <= elif_purpose <= 4
True
```
